# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE

    #raise NotImplementedError()

```

Output:
```
File "<ipython-input-1-ce2cd72dff1d>", line 5
    #raise NotImplementedError()
                                ^
SyntaxError: unexpected EOF while parsing
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

```python
#play with data
import pandas as pd
df = pd.read_csv("assets/NISPUF17.csv", index_col=0)
df.head()
```

Output:
```
SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  AGECPOXR  \
1   128521     12852     2         NaN  235.916956     1031  2017       NaN   
2    10741      1074     2         NaN  957.353840     1068  2017       NaN   
3   220011     22001     2         NaN  189.611299     1050  2017       NaN   
4    86131      8613     1  675.430817  333.447418     1040  2017       NaN   
5   227141     22714     1  482.617748  278.768063     1008  2017       NaN   

   HAD_CPOX  AGEGRP  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  XVRCTY6  \
1         2       1  ...                        NaN      NaN      NaN   
2         2       1  ...                        NaN      NaN      NaN   
3         2       3  ...                        NaN      NaN      NaN   
4         2       1  ...                        NaN      NaN      NaN   
5         2       1  ...                        NaN      NaN      NaN   

   XVRCTY7  XVRCTY8  XVRCTY9  INS_STAT2_I  INS_BREAK_I  
1      NaN      NaN      NaN    
```

```python
import pandas as pd
df=pd.read_csv("assets/NISPUF17.csv", index_col=0)
print(df['EDUC1'])
```

Output:
```
1        4
2        3
3        3
4        4
5        1
        ..
28461    3
28462    2
28463    3
28464    2
28465    4
Name: EDUC1, Length: 28465, dtype: int64
```

```python
Total = len(df['EDUC1'])
print(Total)
```

Output:
```
28465
```

```python
gr1 = len(df[df['EDUC1']==1])
print("Group 1 - :", gr1)

gr2 = len(df[df['EDUC1']==2])
print("Group 2 - :", gr2)

gr3 = len(df[df['EDUC1']==3])
print("Group 3 - :", gr3)

gr4 = len(df[df['EDUC1']==4])
print("Group 4 - :", gr4)
```

Output:
```
Group 1 - : 2904
Group 2 - : 4906
Group 3 - : 6999
Group 4 - : 13656
```

```python
#check if any missing values
diff = gr1 + gr2 + gr3 + gr4 - Total
print(diff)
```

Output:
```
0
```

```python
ratio1 = gr1 / Total
ratio2 = gr2 / Total
ratio3 = gr3 / Total
ratio4 = gr4 / Total
print(ratio1, ratio2, ratio3, ratio4 )
```

Output:
```
0.10202002459160373 0.172352011241876 0.24588090637625154 0.47974705779026877
```

```python
# Calculate proportions by dividing each group by the total
Results = {"less than high school": ratio1,
"high school": ratio2,
"more than high school but not college": ratio3,
"college": ratio4}
print(Results)
```

Output:
```
{'less than high school': 0.10202002459160373, 'high school': 0.172352011241876, 'more than high school but not college': 0.24588090637625154, 'college': 0.47974705779026877}
```

```python
#final submit
def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE
    import pandas as pd
    df=pd.read_csv("assets/NISPUF17.csv", index_col=0)
    total = len(df['EDUC1'])
    gr1 = len(df[df['EDUC1']==1])
    gr2 = len(df[df['EDUC1']==2])
    gr3 = len(df[df['EDUC1']==3])
    gr4 = len(df[df['EDUC1']==4])
    result = {"less than high school":gr1/total,
    "high school":gr2/total,
    "more than high school but not college":gr3/total,
    "college":gr4/total}
    return result
    #raise NotImplementedError()
```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

Output:
```
<ipython-input-11-380749c0010d> in <module>
----> 1 assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

<ipython-input-10-b22e4dde7c56> in average_influenza_doses()
      1 def average_influenza_doses():
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

```python
import pandas as pd
df=pd.read_csv("assets/NISPUF17.csv", index_col=0)
```

```python
df1=df[['CBF_01','P_NUMFLU']]
print(df1)

```

Output:
```
CBF_01  P_NUMFLU
1           1       NaN
2           2       NaN
3           2       NaN
4           2       3.0
5           1       0.0
...       ...       ...
28461       2       NaN
28462       1       NaN
28463       1       NaN
28464       2       NaN
28465       1       NaN

[28465 rows x 2 columns]
```

```python
df1=df[['CBF_01','P_NUMFLU']].dropna()
print(df1)
```

Output:
```
CBF_01  P_NUMFLU
4           2       3.0
5           1       0.0
7           1       0.0
8           1       0.0
9           1       0.0
...       ...       ...
28447       1       2.0
28448       2       1.0
28450       1       4.0
28453       2       2.0
28455       1       0.0

[15333 rows x 2 columns]
```

```python
bm1 = df1[df1['CBF_01']==1].loc[:,'P_NUMFLU']
bm1_av = bm1.mean()
print(bm1_av)
```

Output:
```
1.8799187420058687
```

```python
bm1 = df1[df1['CBF_01']==1].loc[:,'P_NUMFLU'].mean()
print(bm1)
```

Output:
```
1.8799187420058687
```

```python
bm1 = df1[df1['CBF_01']==1].loc[:,'P_NUMFLU'].mean()
print(bm1)
bm2 = df1[df1['CBF_01']==2].loc[:,'P_NUMFLU'].mean()
print(bm2)
print(tuple((bm1,bm2)))
```

Output:
```
1.8799187420058687
1.5963945918878317
(1.8799187420058687, 1.5963945918878317)
```

```python
def average_influenza_doses():
    # YOUR CODE HERE
    import pandas as pd
    df=pd.read_csv("assets/NISPUF17.csv", index_col=0)
    df1=df[['CBF_01','P_NUMFLU']].dropna()
    
    bm1 = df1[df1['CBF_01']==1].loc[:,'P_NUMFLU'].mean()
    bm2 = df1[df1['CBF_01']==2].loc[:,'P_NUMFLU'].mean()
    
    return bm1, bm2
```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

```python
import pandas as pd
df=pd.read_csv("assets/NISPUF17.csv", index_col=0)
print(df[['SEX','HAD_CPOX','P_NUMVRC']])
```

Output:
```
SEX  HAD_CPOX  P_NUMVRC
1        1         2       NaN
2        1         2       NaN
3        2         2       NaN
4        2         2       1.0
5        2         2       0.0
...    ...       ...       ...
28461    2         2       NaN
28462    2         2       NaN
28463    2         2       NaN
28464    2         2       NaN
28465    2         2       NaN

[28465 rows x 3 columns]
```

```python
#males
dfm = df[df['SEX']==1].loc[:,['HAD_CPOX','P_NUMVRC']]
#males ratio
ratiom = len(dfm[ (dfm['HAD_CPOX']==1) & (dfm['P_NUMVRC']>=1)]) / len(dfm[ (dfm['HAD_CPOX']==2) & (dfm['P_NUMVRC']>=1)])
print(ratiom)
```

Output:
```
0.009675583380762664
```

```python
#both ratios
#males
dfm = df[df['SEX']==1].loc[:,['HAD_CPOX','P_NUMVRC']]
#males ratio
ratiom = len(dfm[ (dfm['HAD_CPOX']==1) & (dfm['P_NUMVRC']>=1)]) / len(dfm[ (dfm['HAD_CPOX']==2) & (dfm['P_NUMVRC']>=1)])
print(ratiom)
#females
dff = df[df['SEX']==2].loc[:,['HAD_CPOX','P_NUMVRC']]
#females ratio
ratiof = len(dff[ (dff['HAD_CPOX']==1) & (dff['P_NUMVRC']>=1)]) / len(dff[ (dff['HAD_CPOX']==2) & (dff['P_NUMVRC']>=1)])
print(ratiof)
```

Output:
```
0.009675583380762664
0.0077918259335489565
```

```python
def chickenpox_by_sex():
    import pandas as pd
    df=pd.read_csv("assets/NISPUF17.csv", index_col=0)

    #Males
    dfm = df[df['SEX']==1].loc[:,['HAD_CPOX','P_NUMVRC']]
    ratiom = len(dfm[ (dfm['HAD_CPOX']==1) & (dfm['P_NUMVRC']>=1)]) / len(dfm[ (dfm['HAD_CPOX']==2) & (dfm['P_NUMVRC']>=1)])

    #females
    dff = df[df['SEX']==2].loc[:,['HAD_CPOX','P_NUMVRC']]
    ratiof = len(dff[ (dff['HAD_CPOX']==1) & (dff['P_NUMVRC']>=1)]) / len(dff[ (dff['HAD_CPOX']==2) & (dff['P_NUMVRC']>=1)])

    return {"male":ratiom,"female":ratiof}
```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
                   "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    #return corr

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python
pip install scipy
```

Output:
```
Requirement already satisfied: scipy in c:\users\thoma\anaconda3\lib\site-packages (1.3.1)
Note: you may need to restart the kernel to use updated packages.
```

```python
import scipy.stats as stats
import numpy as np
import pandas as pd

df=pd.read_csv("assets/NISPUF17.csv", index_col=0)
df1 = df[["HAD_CPOX","P_NUMVRC"]].dropna()
```

```python
cor = df1[(df1["HAD_CPOX"]==1) | (df1["HAD_CPOX"]==2)]
print(cor)
```

Output:
```
HAD_CPOX  P_NUMVRC
4             2       1.0
5             2       0.0
7             2       1.0
8             2       0.0
9             1       0.0
...         ...       ...
28447         2       1.0
28448         2       1.0
28450         2       1.0
28453         2       1.0
28455         2       1.0

[15286 rows x 2 columns]
```

```python
corr, pval=stats.pearsonr(cor["HAD_CPOX"],cor["P_NUMVRC"])
print((corr,pval))

```

Output:
```
(0.07044873460148, 2.778026318286582e-18)
```

```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd

    df=pd.read_csv("assets/NISPUF17.csv", index_col=0)
    df1 = df[["HAD_CPOX","P_NUMVRC"]].dropna()

    cor = df1[(df1["HAD_CPOX"]==1) | (df1["HAD_CPOX"]==2)]

    corr, pval=stats.pearsonr(cor["HAD_CPOX"],cor["P_NUMVRC"])

    return corr
```

```python
print(corr)
print(pval)
```

Output:
```
0.07044873460148
2.778026318286582e-18
```

```python

```